```python
# Initialize Otter
import otter
grader = otter.Notebook("cee175.ipynb")
```

# CEE 175: Geotechnical and Geoenvironmental Engineering

> **Mohamad M. Hallal, PhD** <br> Teaching Professor, UC Berkeley

[![License](https://img.shields.io/badge/license-CC%20BY--NC--ND%204.0-blue)](https://creativecommons.org/licenses/by-nc-nd/4.0/)
***

## About this Assignment

The objective of this assignment is to write a Python function that automatically determines the soil classification according to ASTM/USCS standards. You will create several intermediate functions that will ultimately be combined into a single function that outputs the soil classification.

## Instructions

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and these should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

You must submit the assignment to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Autograding and Gradescope

We will be using an automatic grader for this assignment. Each question will have an **Answer Cell** followed by a **Testing Cell**. You will put gradable code into the **Answer Cell** as instructed, and then test your code by executing the **Testing Cell**. The output of the testing cell will let you know whether you have a mistake in the Answer Cell. You can run the Testing Cell infinite times.

### Example Question

You can run individual cells, a selection of cells, or all cells, and there are multiple ways to do so. To run individual cells, you can click on the cell and then:
* From the Toolbar, click the Run icon `▶ Run`
* From the keyboard, press <kbd>Shift</kbd>+<kbd>Enter</kbd>. This shortcut will run the selected cell and *go to the next cell or add a new code cell below if this is the last cell*.
* From the keyboard, press <kbd>Ctrl</kbd>+<kbd>Enter</kbd> (<kbd>Cmd</kbd>+<kbd>Enter</kbd> for macOS). This shortcut will run the selected cell *without going to the next cell*.
* From the Menu bar, click **Run**, then **Run Selected Cell** (there are other run options in the Run menu)

You can run multiple cells from the Menu bar. Click **Run** from the Menu bar. You will see several options for running multiple cells:
* **Run All Cells**: this will run **all cells** in the notebook
* **Run All Above Selected Cell**: this will run **all cells above (but not including)** the selected cell
* **Run All Selected Cell and All Below**: this will run **the selected cell and all cells below it**

Now let's try a simple example. Given two variables, `m=1` and `n=2`, compute their sum and assign it to variable `o`. 

+ The correct solution to this exercise is to replace the `...` below with `m + n`.  Try that out, and confirm that the testing cell marks your answer as correct. 
+ Also try an *incorrect* answer, by replacing the `...` with `-2`. Notice what happens when you run the testing cell.

It is important that you save your work regularly, preferably after answering each question. You can save your work by clicking on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. Failing to save your work may result in incomplete submissions and loss of points.

#### Answer Cell:

```python
m = 1
n = 2
o = ...
```

#### Testing Cell:

```python
grader.check("q0")
```

## Question 1: Plasticity Index 

In programming, a function is a block of code that performs a specific task. Similar to mathematical functions, Python functions take in data, perform on them a sequence of statements, and then return output(s).

An example of a mathematical function is: $f(x)=4x^2+2x-1$

In Python, a function that serves the same task looks like:

```python
def f(x):
    return 4 * x ** 2 + 2 * x - 1
```

The most common way to define a new function in Python is using the keyword `def` (short for define), as shown below:

```python
def function_name(argument_1, argument_2, ...):
    
    # function_body
    
    return output_parameters
```

The syntax includes:
1. Function header that begins with the keyword `def`, which is used to define a function
2. `function_name`, which is the identifier chosen by the programmer (you)
3. Input arguments `(argument_1, argument_2, ...)`, which are a comma-separated sequence of inputs that will be sent to the function 
4. `:` colon, which is used at the end of the function header, after the input arguments
6. `function_body`, which is a sequence of code statements that the function will execute when called
7. `return`, which will return (send back) the value(s) after the keyword to the calling statement

<div class="alert alert-block alert-warning"> <b>NOTE!</b> The code that belongs to a function should be indented, as shown above. Indentation is important as it is the only way Python knows which code belongs to the function. It is standard in Python to use four spaces (equivalent to a tab) for indenting. If the code block is not indented correctly, this will raise an error.</div>

Let's begin my defining our first simple function. Write a function named `plasticity_index()` that calculates and returns the plasticity index (PI) of a soil. This function has the following input arguments:
* `LL`, which is the liquid limit of the soil in %
* `PL`, which is the plastic limit of the soil in %

Please note that variable and function names in Python are case sensitive.

After defining a function, and to use it, we need to **call** the function. To call a function, use its name followed by the values of the input arguments between parentheses: `function_name(argument_1, argument_2, ...)`. For the function that you will define, since the function is named `plasticity_index` and it takes two arguments, you can call it using the following syntax: `plasticity_index(30, 10)`. This will output the value 20, which is the plasticity index based on the given inputs.

Test your `plasticity_index()` function to calculate the plasticity index of a soil with a liquid limit of 55% and a plastic limit of 25%. Assign the output to `q1`. Feel free to test it for any other inputs. Note that when calling the function, you should not include the units such as %. Only include the numeric value inside the parentheses.

```python
# ANSWER CELL
# Only modify ...

def plasticity_index(LL, PL):
    PI = ...
    return PI 
```

```python
# TEST YOUR FUNCTION HERE

q1 = ...

# print result
print(f'PI = {q1} %')
```

```python
grader.check("q1")
```

## Question 2: Coarse- or Fine-Grained?

Functions in Python can be written to do much more than evaluate a mathematical equation. Next, we will write another simple function that tells us whether a soil is coarse- or fine-grained. To do so, we will utilize something known as **conditional statements**.

Every day, we make decisions based on certain conditions. For example, **if** it is raining, I will bring an umbrella before going outside. Similarly, we want computer programs to perform different actions based on specific conditions. This can be achieved using **conditional statements**.

For example, if we want to say something like *if this condition is true, do this, otherwise, do something else*, we can use the `if-else` syntax in Python:

```python
if condition:
    # do code block 1
else:
    # do code block 2
```

Here's how it works:

1. When Python encounters an if statement, it evaluates the specified condition.
2. If the condition is `True`, Python executes code block 1.
3. If the condition is `False`, Python skips code block 1 and proceeds to the else statement, executing code block 2.

Here's a breakdown of the syntax:

* The keyword `if` is followed by the condition, which is an expression that evaluates to either `True` or `False`.
* A colon `:` follows the condition, indicating the start of the code block.
* The code block, which will only be executed if the condition is `True`, follows the colon. This block can contain one or more lines of code.
* The code block *must be indented*. Indentation is crucial because it indicates which code belongs to the if statement.
* The keyword `else` is used to define an alternate code block that will execute only if the condition in the `if` statement evaluates to `False`.
* Like the `if` statement, the `else` statement ends with a colon `:` to indicate the start of the alternate code block.
* The code block associated with `else` must also be indented, just like the `if` block. This indentation ensures that Python knows which code belongs to the `else` statement.
 
Write a function named `coarse()` that checks whether a soil is coarse- or fine-grained. This function has the following input argument:
* `P200`, which is the percentage of soil passing through Sieve #200 (i.e., percent finer) by weight (or mass)

Your function should return the integer `1` if the soil is coarse-grained and the integer `0` otherwise.

Test your `coarse()` function to check if a soil with `P200` equal to 40% is coarse- or fine-grained. Assign the output to `q2`. Feel free to test it for any other inputs.

```python
# ANSWER CELL

def coarse(P200):

    if ...:
        return ...
    else:
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q2 = ...

# print result
print(f'The soil is coarse-grained: {bool(q2)}')
```

```python
grader.check("q2")
```

## Question 3: Gravel or Sand?

Assume that a soil is coarse-grained. We want to check if it is gravel or sand.
 
Write a function named `gravel_or_sand()` that checks whether a coarse-grained soil is gravel or sand. This function has the following input arguments:
* `P4`, which is the percentage of soil passing through Sieve #4 (i.e., percent finer) by weight (or mass)
* `P200`, which is the percentage of soil passing through Sieve #200 (i.e., percent finer) by weight (or mass)

Your function should return the string `'G'` if the soil is gravel and the string `'S'` if it is sand. Note that strings are a sequence of characters enclosed within quotes, such as `' '`, so don't forget to include the quotes!

Test your `gravel_or_sand()` function to check if a soil with `P4` equal to 70% and `P200` equal to 45% is gravel or sand. Assign the output to `q3`. Feel free to test it for any other inputs from the ones shown below.

```PYTHON
Examples:

>>> gravel_or_sand(70, 45)
'G'

>>> gravel_or_sand(70, 35)
'S'

>>> gravel_or_sand(70, 40)
'S'
```

```python
# ANSWER CELL

def gravel_or_sand(P4, P200):

    # compute the coarse fraction of the soil
    coarse_fraction = ...

    # compute the gravel fraction of the soil
    gravel_fraction = ...

    # check if gravel or sand
    if ...:
        return ...
    else:
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q3 = ...

# print result
print(f"The soil is {'gravel' if q3.upper()=='G' else 'sand' if q3.upper()=='S' else None}.")
```

```python
grader.check("q3")
```

## Question 4: Coefficients of Uniformity and Curvature

Assume that a soil is coarse-grained. We want to calculate its coefficients of uniformity and curvature, $Cu$ and $Cc$, respectively.
 
Write a function named `coefficients()` that calculates the coefficients of uniformity and curvature of a soil. This function has the following input arguments:
* `D60`, which is the particle size corresponding to 60% finer in mm
* `D30`, which is the particle size corresponding to 30% finer in mm
* `D10`, which is the particle size corresponding to 10% finer in mm

Your function should compute and return the following (in this order):
1. `Cu`: The coefficient of uniformity
2. `Cc`: The coefficient of curvature

Note that Python uses the following operators/symbols to perform different mathematical operations.

| Operation       | Mathematical Notation | Python Operator | Python Example | Output    |
|:----------------|:---------------------:|:---------------:|:---------------|:----------|
| Addition        | $a+b$                 | `+`             | `3 + 2`        | `5`       |
| Subtraction     | $a-b$                 | `-`             | `3 - 2`        | `1`       |
| Multiplication  | $a\times b$           | `*`             | `3 * 2`        | `6`       |
| Division        | $\dfrac{a}{b}$        | `/`             | `3 / 2`        | `1.5`     |
| Exponentiation  | $a^b$                 | `**`            | `3 ** 2`       | `9`       |

Test your `coefficients()` function using a `D60` equal to 0.15 mm, `D30` equal to 0.13 mm, and `D10` equal to 0.1 mm and assign the outputs to `Cu` and `Cc`, respectively. Feel free to test it for any other inputs.

```python
# ANSWER CELL

def coefficients(D60, D30, D10):

    Cu = ...
    Cc = ...
    
    return Cu, Cc
```

```python
# TEST YOUR FUNCTION HERE

Cu, Cc = ...

# print results
print(f'Cu = {Cu}\nCc = {Cc}')
```

```python
grader.check("q4")
```

## Question 5: Well or Poorly Graded?

Assume that a soil is coarse-grained. We want to check if it is well or poorly graded.
 
Write a function named `gradation()` that checks if a coarse-grained soil is well or poorly graded. This function has the following input arguments:
* `P4`, which is the percentage of soil passing through Sieve #4 (i.e., percent finer) by weight (or mass)
* `P200`, which is the percentage of soil passing through Sieve #200 (i.e., percent finer) by weight (or mass)
* `D60`, which is the particle size corresponding to 60% finer in mm
* `D30`, which is the particle size corresponding to 30% finer in mm
* `D10`, which is the particle size corresponding to 10% finer in mm

Your function should return the string `'W'` if the soil is well-graded and the string `'P'` if it is poorly graded.

Since the conditions for well-graded and poorly graded soil differ between gravel and sand, you must first determine whether the soil is gravel or sand and then apply the corresponding conditions for each. To do this, use the `gravel_or_sand()` function you previously wrote, which checks whether a coarse-grained soil is gravel or sand. You can call this function within the implementation of the function below.

To check multiple conditions in Python, you can use logical operators such as `and`. For example, if you want to check if a variable `var1` is greater than 3 and another variable `var2` is less than 2, you can combine the conditions using the `and` operator like this:

``` python
var1 > 3 and var2 < 2
```

This expression will evaluate to True **only if both conditions** are true. If either condition is false, it will evaluate to False. Python includes a variety of comparison operators for performing comparisons:

| Comparison         | Python Operator | True example | False Example |
|:-------------------|:---------------:|:-------------|:--------------|
| Less than          | `<`             | `2 < 3`      | `2 < 2`       |
| Greater than       | `>`             | `3 > 2`      | `3 > 3`       |
| Less than or equal | `<=`            | `2 <= 2`     | `3 <= 2`      |
| Greater or equal   | `>=`            | `3 >= 3`     | `2 >= 3`      |
| Equal              | `==`            | `3 == 3`     | `3 == 2`      |
| Not equal          | `!=`            | `3 != 2`     | `2 != 2`      |

<div class="alert alert-block alert-warning"> <b>NOTE!</b> Python allows string comparison. So, if you want to check if the output of the fucntion <code>coarse()</code> is gravel, you can use the following: <br> <code>coarse(P4, P200) == 'G'</code>. This will evaluate to True if the soil is gravel and False otherwise. Note that, as shown in the table above, Python uses a double equal sign <code>==</code> to check if two expressions are equal to one another.</div>

Test your `gradation()` function using a soil with `P4` equal to 70%, `P200` equal to 45%, `D60` equal to 0.15 mm, `D30` equal to 0.13 mm, and `D10` equal to 0.1 mm and assign the output to `q5`. Feel free to test it for any other inputs from the ones shown below.

```PYTHON
Examples:

>>> gradation(70, 45, 0.15, 0.13, 0.1)
'P'

>>> gradation(70, 45, 0.45, 0.3, 0.1)
'W'

>>> gradation(70, 35, 0.15, 0.14, 0.005)
'P'

>>> gradation(70, 35, 0.65, 0.3, 0.1)
'W'
```

Note that the function below uses **nested if statement**, which is an `if` statement inside another `if` statement. This allows for more complex conditions and actions. The syntax for nested `if` statements is:

```python
if condition_1: # outer if 
    
    if condition_2: # inner if 
        # do code block 1a
    else: # inner else 
        # do code block 1b
        
else: # outer else
    
    if condition_3: # inner if 
        # do code block 2a
    else: # inner else 
        # do code block 2b
```

Because the inner `if` statement is nested within the outer `if` statement, both condition_1 and condition_2 should be `True` for the innermost code block 1a to execute. This means that even if condition_2 is `True` but condition_1 is `False`, code block 1a will not be executed. In addition, the inner `if` statement must be indented further than the outer `if` statement to indicate that it is part of the outer `if` block. Proper indentation is crucial for readability and correct program execution.

```python
# ANSWER CELL

def gradation(P4, P200, D60, D30, D10):

    # call the function you defined to compute Cu and Cc
    Cu, Cc = coefficients(D60, D30, D10)
    
    if ...: # add condition that evaluates to true if the soil is gravel by calling your gravel_or_sand() function
        if ...: # add condition that evaluates to true if the gravel is well-graded
            return ...
        else: # otherwise, the gravel is poorly graded
            return ...
    else: # otherwise, the soil is sand
        if ...: # add condition that evaluates to true if the sand is well-graded
            return ...
        else: # otherwise, the sand is poorly graded
            return ...
```

```python
# TEST YOUR FUNCTION HERE

q5 = ...

# print result
print(f"The soil is {'well-graded' if q5.upper()=='W' else 'poorly graded' if q5.upper()=='P' else None}.")
```

```python
grader.check("q5")
```

The functions that you have written thus far are mostly focused on coarse-grained soils. Next, we will focus on writing functions to help with classifying fine-grained soils. Finally, we will combine all functions together to develop a complete soil classification function. 

## Question 6: Low or High Plasticity?

Fine-grained soils are mainly classified based on whether they have low or high plasticity and where they fall with respect to the "A"-line.
 
Write a function named `plasticity()` that checks if a fine-grained soil or the fine-grained portion of a soil has low or high plasticity. This function has the following input argument:
* `LL`, which is the liquid limit of the soil in %

Your function should return the string `'L'` if the soil has low plasticity and the string `'H'` if it has high plasticity.

Test your `plasticity()` function using a soil with `LL` equal to 55% and assign the output to `q6`. Feel free to test it for any other inputs.

Note that for this question, we are providing you with a blank cell. By now, you should have seen several examples on how to define a function and structure it. So for this question, you should:
1. Define the function and include its name and input arguments
2. Include the body of the function that checks the conditions and returns the appropriate output
3. Don't forget to include colons `:` where required. Mainly, you should include a `:` after defining a function and at the end of a conditional statement (`if`, `else`, etc.).
4. Pay attention to indentation
5. You may refer to previous examples

```python
# ANSWER CELL

...
```

```python
# TEST YOUR FUNCTION HERE

q6 = ...

# print result
print(f"The soil has {'low plasticity' if q6.upper()=='L' else 'high plasticity' if q6.upper()=='H' else None}.")
```

```python
grader.check("q6")
```

## Question 7: Above or Below "A"-Line?

Next, let's classify the fines based on where the plasticity index falls with respect to the "A"-line.
 
Write a function named `silt_or_clay()` that checks if a fine-grained soil or the fine-grained portion of a soil is above or below the "A"-line. This function has the following input arguments:
* `LL`, which is the liquid limit of the soil in %
* `PL`, which is the plastic limit of the soil in %

Your function should return the string `'C'` if the plasticity index of the soil plots on or above the "A"-line and the string `'M'` otherwise.

Since you must first determine the plasticity index of the soil, use the `plasticity_index()` function you previously defined. You can call this function within the implementation of the function below.

Note that the "A"-line has the following equations:

$$A_{PI} =
\begin{cases} 
4, & \text{if } LL \leq 25.5 \%\\
0.73 (LL - 20), & \text{if } LL > 25.5 \%
\end{cases}$$

where $A_{PI}$ is the plasticity index exactly on the "A"-line.

Test your `silt_or_clay()` function using a soil with a liquid limit of 55% and a plastic limit of 25% and assign the output to `q7`. Feel free to test it for any other inputs from the ones shown below.

```PYTHON
Examples:

>>> silt_or_clay(55, 25)
'C'

>>> silt_or_clay(55, 45)
'M'

>>> silt_or_clay(10, 7)
'M'

>>> silt_or_clay(20, 10)
'C'
```

```python
# ANSWER CELL

def silt_or_clay(LL, PL):

    # call the plasticity_index() function you defined to compute the actual PI
    actual_PI = ...

    # get PI on the "A"-line
    if ...:
        A_PI = ...
    else:
        A_PI = ...

    # compare actual PI and that on the "A"-line and return the correct output 
    if ...:
        return ...
    else:
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q7 = ...

# print result
print(f"The soil is {'clay' if q7.upper()=='C' else 'silt' if q7.upper()=='M' else None}.")
```

```python
grader.check("q7")
```

Now, you should have all the intermediate functions needed to classify both coarse- and fine-grained soils. Next, we will write a function to determine the group symbol for fine-grained soils and then we'll do the same for coarse-grained soils.

## Question 8: Group Symbol for Fine-Grained Soils

For simplicity, we will only focus on inorganic soils.

Write a function named `classify_fine()` that returns the group symbol classification for a fine-grained soil or the fine-grained portion of a soil. This function has the following input arguments:
* `LL`, which is the liquid limit of the soil in %
* `PL`, which is the plastic limit of the soil in %

Your function should return one of the following strings `'CL'`, `'ML'`, `'CH'`, `'MH'`, or `'CL-ML'`. Refer to the lecture material for the requirements of each group symbol.

Since you must first determine whether the soil has low or high plasticity and whether the plasticity index plots above or below the "A"-line, use the `plasticity()` and `silt_or_clay()` functions you previously defined. You can call these functions within the implementation of the function below.

***Hint**: It should be obvious that `'CL'`, `'ML'`, `'CH'`, and `'MH'` are a combination of the outputs from the `silt_or_clay()` and `plasticity()` functions. Meaning, if `silt_or_clay(55, 25)` returns `'C'` and `plasticity(55)` returns `'H'`, the group symbol is `'CH'` (combination of the two outputs). The only exception to this rule is the `'CL-ML'` group symbol. So here is how we suggest you approach this question:*
1. Check the condition for `'CL-ML'`, and if satisfied (i.e., True), return `'CL-ML'` (refer to the lecture notes for the condition)
2. Else, if the condition for `'CL-ML'` is not satisfied (i.e., False), return the combination of the outputs from the `silt_or_clay()` and `plasticity()` functions

An easy way to combine the outputs of the `silt_or_clay()` and `plasticity()` functions is by placing a plus sign `+` between the function calls.  This will concatenate or "glue" their outputs together into a single string. Concatenation is a simple way to merge multiple strings into one. 

For example, `silt_or_clay(55, 25) + plasticity(55)` will evaluate to `'CH'` (i.e., `'C' + 'H'`).


Test your `classify_fine()` function using a soil with a liquid limit of 55% and a plastic limit of 25% and assign the output to `q8`. Feel free to test it for any other inputs from the ones shown below.

```PYTHON
Examples:

>>> classify_fine(55, 25)
'CH'

>>> classify_fine(55, 45)
'MH'

>>> classify_fine(10, 7)
'ML'

>>> classify_fine(20, 10)
'CL'

>>> classify_fine(10, 5)
'CL-ML'
```

```python
# ANSWER CELL

def classify_fine(LL, PL):
    
    if ...:
        return ...
    else:
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q8 = ...

# print result
print(f'The soil is {q8}.')
```

```python
grader.check("q8")
```

## Question 9: Group Symbol for Gravel

Write a function named `classify_gravel()` that returns the group symbol classification for gravel. This function has the following input arguments:
* `P4`, which is the percentage of soil passing through Sieve #4 (i.e., percent finer) by weight (or mass)
* `P200`, which is the percentage of soil passing through Sieve #200 (i.e., percent finer) by weight (or mass)
* `D60`, which is the particle size corresponding to 60% finer in mm
* `D30`, which is the particle size corresponding to 30% finer in mm
* `D10`, which is the particle size corresponding to 10% finer in mm
* `LL`, which is the liquid limit of the soil in %
* `PL`, which is the plastic limit of the soil in %

Your function should return one of the following strings `'GW'`, `'GP'`, `'GW-GM'`, `'GW-GC'`, `'GP-GM'`, `'GP-GC'`, `'GM'`, `'GC'`, or `'GC-GM'`. Refer to the lecture material for the requirements of each group symbol.

Since you must first determine whether the soil is well- or poorly graded and in some cases the classification of the fine-grained portion of the soil, use the `gradation()` and `classify_fine()` functions you previously defined. You can call these functions within the implementation of the function below.

If a soil does not have LL and PL, assign them values of 0 when calling the function.

Test your `classify_gravel()` function using any of the examples below and assign the output to `q9`. Feel free to test it for any other inputs from the ones shown below.

```PYTHON
Examples:

>>> classify_gravel(40, 4, 0.3, 0.14, 0.025, 0, 0)
'GW'

>>> classify_gravel(40, 4, 0.15, 0.13, 0.1, 0, 0)
'GP'

>>> classify_gravel(40, 10, 0.3, 0.14, 0.025, 30, 25)
'GW-GM'

>>> classify_gravel(40, 10, 0.3, 0.14, 0.025, 30, 15)
'GW-GC'

>>> classify_gravel(40, 10, 0.15, 0.13, 0.1, 30, 25)
'GP-GM'

>>> classify_gravel(40, 10, 0.15, 0.13, 0.1, 30, 15)
'GP-GC'

>>> classify_gravel(40, 20, 0.15, 0.13, 0.1, 30, 25)
'GM'

>>> classify_gravel(40, 20, 0.15, 0.13, 0.1, 30, 15)
'GC'

>>> classify_gravel(40, 20, 0.15, 0.13, 0.1, 25, 20)
'GC-GM'
```

```python
# ANSWER CELL

def classify_gravel(P4, P200, D60, D30, D10, LL, PL):

    # Get the gradation result and fine classification
    gradation_result = ...
    fine_group = ...
    
    # Case 1: Very little fines (< 5%)
    if ...:
        if ...:  # Well-graded
            return ...
        else:  # Poorly graded
            return ...

    # Case 2: Moderate fines (5-12% fines)
    if ...:
        if ...:  # Well-graded
            if ...:  # Silts
                return ...
            else:  # Clays
                return ...
        else:  # Poorly graded
            if ...:  # Silts
                return ...
            else:  # Clays
                return ...

    # Case 3: High fines (> 12%)
    # There is no need to check if fines greater than 12% 
    # The code will only reach this point if the two previous conditions (< 5% and 5-12%) are False
    if ...:  # Silts
        return ...
    elif ...:  # Clays
        return ...
    else:  # Mixture of fines
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q9 = ...

# print result
print(f'The soil is {q9}.')
```

```python
grader.check("q9")
```

## Question 10: Group Symbol for Sand

Write a function named `classify_sand()` that returns the group symbol classification for sand. This function has the following input arguments:
* `P4`, which is the percentage of soil passing through Sieve #4 (i.e., percent finer) by weight (or mass)
* `P200`, which is the percentage of soil passing through Sieve #200 (i.e., percent finer) by weight (or mass)
* `D60`, which is the particle size corresponding to 60% finer in mm
* `D30`, which is the particle size corresponding to 30% finer in mm
* `D10`, which is the particle size corresponding to 10% finer in mm
* `LL`, which is the liquid limit of the soil in %
* `PL`, which is the plastic limit of the soil in %

Your function should return one of the following strings `'SW'`, `'SP'`, `'SW-SM'`, `'SW-SC'`, `'SP-SM'`, `'SP-SC'`, `'SM'`, `'SC'`, or `'SC-SM'`. Refer to the lecture material for the requirements of each group symbol.

Since you must first determine whether the soil is well- or poorly graded and in some cases the classification of the fine-grained portion of the soil, use the `gradation()` and `classify_fine()` functions you previously defined. You can call these functions within the implementation of the function below.

If a soil does not have LL and PL, assign them values of 0 when calling the function.

**Note** that if you implemented Question 9 correctly, you should be able to copy its code and paste below and simply replace 'G' with 'S'.

Test your `classify_sand()` function using any of the examples below and assign the output to `q10`. Feel free to test it for any other inputs from the ones shown below.

```PYTHON
Examples:

>>> classify_sand(60, 4, 0.3, 0.14, 0.025, 0, 0)
'SW'

>>> classify_sand(60, 4, 0.15, 0.13, 0.1, 0, 0)
'SP'

>>> classify_sand(60, 10, 0.3, 0.14, 0.025, 30, 25)
'SW-SM'

>>> classify_sand(60, 10, 0.3, 0.14, 0.025, 30, 15)
'SW-SC'

>>> classify_sand(60, 10, 0.15, 0.13, 0.1, 30, 25)
'SP-SM'

>>> classify_sand(60, 10, 0.15, 0.13, 0.1, 30, 15)
'SP-SC'

>>> classify_sand(65, 20, 0.15, 0.13, 0.1, 30, 25)
'SM'

>>> classify_sand(65, 20, 0.15, 0.13, 0.1, 30, 15)
'SC'

>>> classify_sand(65, 20, 0.15, 0.13, 0.1, 25, 20)
'SC-SM'
```

```python
# ANSWER CELL

def classify_sand(P4, P200, D60, D30, D10, LL, PL):

    # Get the gradation result and fine classification
    gradation_result = ...
    fine_group = ...
    
    # Case 1: Very little fines (< 5%)
    if ...:
        if ...:  # Well-graded
            return ...
        else:  # Poorly graded
            return ...

    # Case 2: Moderate fines (5-12% fines)
    if ...:
        if ...:  # Well-graded
            if ...:  # Silts
                return ...
            else:  # Clays
                return ...
        else:  # Poorly graded
            if ...:  # Silts
                return ...
            else:  # Clays
                return ...

    # Case 3: High fines (> 12%)
    # There is no need to check if fines greater than 12% 
    # The code will only reach this point if the two previous conditions (< 5% and 5-12%) are False
    if ...:  # Silts
        return ...
    elif ...:  # Clays
        return ...
    else:  # Mixture of fines
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q10 = ...

# print result
print(f'The soil is {q10}.')
```

```python
grader.check("q10")
```

## Question 11: Full Soil Classification

Now that we have all the pieces, let's put them together to write one main function that determines the soil classification of any soil.

Write a function named `classify_soil()` that returns the group symbol classification for any soil. This function has the following input arguments:
* `P4`, which is the percentage of soil passing through Sieve #4 (i.e., percent finer) by weight (or mass)
* `P200`, which is the percentage of soil passing through Sieve #200 (i.e., percent finer) by weight (or mass)
* `D60`, which is the particle size corresponding to 60% finer in mm
* `D30`, which is the particle size corresponding to 30% finer in mm
* `D10`, which is the particle size corresponding to 10% finer in mm
* `LL`, which is the liquid limit of the soil in %
* `PL`, which is the plastic limit of the soil in %

Here are the steps:
1. Check if the soil is coarse- or fine-grained using the `coarse()` function
2. If coarse-grained, check if it is gravel or sand using the `gravel_or_sand()` function
    1. If gravel, classify it using the `classify_gravel()` function
    2. Else if the soil is sand, classify it using the `classify_sand()` function
3. Else if the soil is fine-grained, classify it using the `classify_fine()` function

Test your `classify_soil()` function using any of the examples below and assign the output to `q11`. Feel free to test it for any other inputs from the ones shown below.

Note that for fine-grained soils, `D60`, `D30`, and `D10` are not required. So you can assign them values of 0 when calling the function.

```PYTHON
Examples:

>>> classify_soil(100, 2.5, 0.15, 0.13, 0.095, 0, 0)
'SP'

>>> classify_soil(75, 7.5, 0.41, 0.23, 0.095, 0, 0)
'SP-SM'

>>> classify_soil(100, 82, 0, 0, 0, 46, 30)
'ML'

>>> classify_soil(95, 54, 0, 0, 0, 75, 32)
'CH'
```

```python
# ANSWER CELL

def classify_soil(P4, P200, D60, D30, D10, LL, PL):
    
    # Check if coarse grained
    if ...:
        if ...:  # Check if gravel
            return ... 
        else: # otherwise, it is sand
            return ...

    else: # otherwise, it is fine grained
        return ...
```

```python
# TEST YOUR FUNCTION HERE

q11 = ...

# print result
print(f'The soil is {q11}.')
```

```python
grader.check("q11")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q0 (points: 0)
```
>>> assert get_hash(o) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q1 (points: None)
```
>>> assert get_hash(type(q1)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(int(plasticity_index(55, 25))) == '34173cb38f07f89ddbebc2ac9128303f'
>>> assert get_hash(int(plasticity_index(60, 40))) == '98f13708210194c475687be6106a3b84'
>>> assert get_hash(int(plasticity_index(30, 5))) == '8e296a067a37563370ded05f5a3bf3ec'
```


Tests for q10 (points: None)
```
>>> assert get_hash(type(q10)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert classify_sand(60, 4, 0.3, 0.14, 0.025, 0, 0).upper() == 'SW'
>>> assert classify_sand(60, 4, 0.15, 0.13, 0.1, 0, 0).upper() == 'SP'
>>> assert classify_sand(60, 10, 0.3, 0.14, 0.025, 30, 25).upper() == 'SW-SM'
>>> assert classify_sand(60, 10, 0.3, 0.14, 0.025, 30, 15).upper() == 'SW-SC'
>>> assert classify_sand(60, 10, 0.15, 0.13, 0.1, 30, 25).upper() == 'SP-SM'
>>> assert classify_sand(60, 10, 0.15, 0.13, 0.1, 30, 15).upper() == 'SP-SC'
>>> assert classify_sand(65, 20, 0.15, 0.13, 0.1, 30, 25).upper() == 'SM'
>>> assert classify_sand(65, 20, 0.15, 0.13, 0.1, 30, 15).upper() == 'SC'
>>> assert classify_sand(65, 20, 0.15, 0.13, 0.1, 25, 20).upper() == 'SC-SM'
```
```
>>> assert get_hash(classify_sand(60, 1, 2, 1.1, 0.3, 0, 0).upper()) == '6f56aa4e2561eb66f17f6d8de8070a77'
>>> assert get_hash(classify_sand(60, 1, 2, 1.5, 0.3, 0, 0).upper()) == '674769e3326f8cf937af4282f2815c02'
>>> assert get_hash(classify_sand(60, 6, 0.7, 0.3, 0.1, 60, 40).upper()) == '65c56b049b73002cf51b3af4e2bd53fe'
>>> assert get_hash(classify_sand(60, 6, 0.7, 0.3, 0.1, 60, 30).upper()) == 'f419abeb5493c8c55303eb149514d422'
>>> assert get_hash(classify_sand(60, 6, 0.45, 0.25, 0.1, 60, 40).upper()) == '1f6c58153add4babee7e71c6901c7134'
>>> assert get_hash(classify_sand(60, 6, 0.45, 0.2, 0.1, 60, 30).upper()) == 'b04d57f315980016bd43ba0ef4a25a3e'
>>> assert get_hash(classify_sand(70, 13, 0.15, 0.13, 0.1, 20, 17).upper()) == '4e0d4f6ce30646f5a3f3e2a7422c1c5a'
>>> assert get_hash(classify_sand(70, 13, 0.15, 0.13, 0.1, 55, 25).upper()) == '6a65edb0cc17d66c677814115b1477f5'
>>> assert get_hash(classify_sand(70, 13, 0.15, 0.13, 0.1, 10, 4).upper()) == 'ed68e6ebff520017f868f48f2dc722c5'
```


Tests for q11 (points: None)
```
>>> assert get_hash(type(q11)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(classify_soil(10, 55, 0, 0, 0, 55, 25).upper()) == '1ee0bf89c5d1032317d13a2e022793c8'
>>> assert get_hash(classify_soil(10, 60, 0, 0, 0, 75, 25).upper()) == '1ee0bf89c5d1032317d13a2e022793c8'
>>> assert get_hash(classify_soil(10, 75, 0, 0, 0, 30, 20).upper()) == '5bc574a47246f122016869b32a6aa6f0'
>>> assert get_hash(classify_soil(10, 60, 0, 0, 0, 60, 30.7).upper()) == '1ee0bf89c5d1032317d13a2e022793c8'
>>> assert get_hash(classify_soil(10, 75, 0, 0, 0, 55, 45).upper()) == '002f27e5064e874ecf4f5def17d1b797'
>>> assert get_hash(classify_soil(10, 80, 0, 0, 0, 75, 45).upper()) == '002f27e5064e874ecf4f5def17d1b797'
>>> assert get_hash(classify_soil(10, 62, 0, 0, 0, 30, 25).upper()) == 'd01fd9b01e9dde8bd3dc247afbfb7218'
>>> assert get_hash(classify_soil(10, 75, 0, 0, 0, 60, 30.9).upper()) == '002f27e5064e874ecf4f5def17d1b797'
```
```
>>> assert get_hash(classify_soil(10, 75, 0, 0, 0, 20, 10).upper()) == '5bc574a47246f122016869b32a6aa6f0'
>>> assert get_hash(classify_soil(10, 67, 0, 0, 0, 10, 7).upper()) == 'd01fd9b01e9dde8bd3dc247afbfb7218'
>>> assert get_hash(classify_soil(10, 59, 0, 0, 0, 20, 17).upper()) == 'd01fd9b01e9dde8bd3dc247afbfb7218'
```
```
>>> assert classify_soil(40, 4, 0.3, 0.14, 0.025, 0, 0).upper() == 'GW'
>>> assert classify_soil(40, 4, 0.15, 0.13, 0.1, 0, 0).upper() == 'GP'
>>> assert classify_soil(40, 10, 0.3, 0.14, 0.025, 30, 25).upper() == 'GW-GM'
>>> assert classify_soil(40, 10, 0.3, 0.14, 0.025, 30, 15).upper() == 'GW-GC'
>>> assert classify_soil(40, 10, 0.15, 0.13, 0.1, 30, 25).upper() == 'GP-GM'
>>> assert classify_soil(40, 10, 0.15, 0.13, 0.1, 30, 15).upper() == 'GP-GC'
>>> assert classify_soil(40, 20, 0.15, 0.13, 0.1, 30, 25).upper() == 'GM'
>>> assert classify_soil(40, 20, 0.15, 0.13, 0.1, 30, 15).upper() == 'GC'
>>> assert classify_soil(40, 20, 0.15, 0.13, 0.1, 25, 20).upper() == 'GC-GM'
```
```
>>> assert get_hash(classify_soil(50, 1, 2, 1.1, 0.3, 0, 0).upper()) == 'c17d19f7520be36addbeb5d2c76ab7cf'
>>> assert get_hash(classify_soil(50, 1, 2, 1.5, 0.3, 0, 0).upper()) == 'ad2d8ee7d788dcf41f399818f639cb64'
>>> assert get_hash(classify_soil(50, 6, 0.45, 0.25, 0.1, 60, 40).upper()) == 'c31f61bf166e30ec55828503fad107d4'
>>> assert get_hash(classify_soil(50, 6, 0.45, 0.25, 0.1, 60, 30).upper()) == 'acb50b1138e657e40d297035097d0bf1'
>>> assert get_hash(classify_soil(50, 6, 0.45, 0.2, 0.1, 60, 40).upper()) == '03f9d3f62bd6324143e47935f642177a'
>>> assert get_hash(classify_soil(50, 6, 0.45, 0.2, 0.1, 60, 30).upper()) == 'f9bb4d123b6b669ac79bd9465d619d00'
>>> assert get_hash(classify_soil(50, 13, 0.15, 0.13, 0.1, 20, 17).upper()) == '64f3bd1741ab8d6ba545a1ae09bb8728'
>>> assert get_hash(classify_soil(50, 13, 0.15, 0.13, 0.1, 55, 25).upper()) == '7359c3eb5c57547295a76ac1bf775b29'
>>> assert get_hash(classify_soil(50, 13, 0.15, 0.13, 0.1, 10, 4).upper()) == '58a3dd33c44967a9f632495eaf1b41b0'
```
```
>>> assert classify_soil(60, 4, 0.3, 0.14, 0.025, 0, 0).upper() == 'SW'
>>> assert classify_soil(60, 4, 0.15, 0.13, 0.1, 0, 0).upper() == 'SP'
>>> assert classify_soil(60, 10, 0.3, 0.14, 0.025, 30, 25).upper() == 'SW-SM'
>>> assert classify_soil(60, 10, 0.3, 0.14, 0.025, 30, 15).upper() == 'SW-SC'
>>> assert classify_soil(60, 10, 0.15, 0.13, 0.1, 30, 25).upper() == 'SP-SM'
>>> assert classify_soil(60, 10, 0.15, 0.13, 0.1, 30, 15).upper() == 'SP-SC'
>>> assert classify_soil(65, 20, 0.15, 0.13, 0.1, 30, 25).upper() == 'SM'
>>> assert classify_soil(65, 20, 0.15, 0.13, 0.1, 30, 15).upper() == 'SC'
>>> assert classify_soil(65, 20, 0.15, 0.13, 0.1, 25, 20).upper() == 'SC-SM'
```
```
>>> assert get_hash(classify_soil(60, 1, 2, 1.1, 0.3, 0, 0).upper()) == '6f56aa4e2561eb66f17f6d8de8070a77'
>>> assert get_hash(classify_soil(60, 1, 2, 1.5, 0.3, 0, 0).upper()) == '674769e3326f8cf937af4282f2815c02'
>>> assert get_hash(classify_soil(60, 6, 0.7, 0.3, 0.1, 60, 40).upper()) == '65c56b049b73002cf51b3af4e2bd53fe'
>>> assert get_hash(classify_soil(60, 6, 0.7, 0.3, 0.1, 60, 30).upper()) == 'f419abeb5493c8c55303eb149514d422'
>>> assert get_hash(classify_soil(60, 6, 0.45, 0.25, 0.1, 60, 40).upper()) == '1f6c58153add4babee7e71c6901c7134'
>>> assert get_hash(classify_soil(60, 6, 0.45, 0.2, 0.1, 60, 30).upper()) == 'b04d57f315980016bd43ba0ef4a25a3e'
>>> assert get_hash(classify_soil(70, 13, 0.15, 0.13, 0.1, 20, 17).upper()) == '4e0d4f6ce30646f5a3f3e2a7422c1c5a'
>>> assert get_hash(classify_soil(70, 13, 0.15, 0.13, 0.1, 55, 25).upper()) == '6a65edb0cc17d66c677814115b1477f5'
>>> assert get_hash(classify_soil(70, 13, 0.15, 0.13, 0.1, 10, 4).upper()) == 'ed68e6ebff520017f868f48f2dc722c5'
```


Tests for q2 (points: None)
```
>>> assert get_hash(type(q2)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(int(coarse(5))) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(int(coarse(40))) == 'c4ca4238a0b923820dcc509a6f75849b'
>>> assert get_hash(int(coarse(55))) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(int(coarse(55))) == 'cfcd208495d565ef66e7dff9f98764da'
```
```
>>> assert get_hash(int(coarse(50))) == 'cfcd208495d565ef66e7dff9f98764da'
```
failure_message: If P200 is exactly equal to 50%, and per USCS, is the soil coarse- or fine-grained?


Tests for q3 (points: None)
```
>>> assert get_hash(type(q3)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(gravel_or_sand(70, 45).upper()) == 'dfcf28d0734569a6a693bc8194de62bf'
>>> assert get_hash(gravel_or_sand(80, 90).upper()) == 'dfcf28d0734569a6a693bc8194de62bf'
>>> assert get_hash(gravel_or_sand(70, 35).upper()) == '5dbc98dcc983a70728bd082d1a47546e'
>>> assert get_hash(gravel_or_sand(80, 50).upper()) == '5dbc98dcc983a70728bd082d1a47546e'
```
```
>>> assert get_hash(gravel_or_sand(70, 40).upper()) == '5dbc98dcc983a70728bd082d1a47546e'
```
failure_message: If exactly 50% of coarse fraction passes sieve


Tests for q4 (points: None)
```
>>> assert get_hash(type(Cu)) != '14e736438b115821cbb9b7ac0ba79034'
>>> assert get_hash(type(Cc)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(round(coefficients(0.15, 0.13, 0.1)[0], 3)) == '6008647277c4454cecd97d33c069f0ca'
>>> assert get_hash(round(coefficients(0.15, 0.13, 0.1)[1], 3)) == 'e1efa1216f20c79a077a45c36dd0bd32'
>>> assert get_hash(round(coefficients(2, 1.1, 0.3)[0], 3)) == 'c3e5bc1a469d73f843b6bcb150cd868e'
>>> assert get_hash(round(coefficients(2, 1.1, 0.3)[1], 3)) == '31ca5cbbaff2e53df4c500d50e94f943'
>>> assert get_hash(round(coefficients(1.1, 0.075, 0.005)[0], 3)) == '2d22f62e9e38119403851fbf17cb26b1'
>>> assert get_hash(round(coefficients(1.1, 0.075, 0.005)[1], 3)) == '4e421203259d2f9b4cad729e5ee3bfd1'
```


Tests for q5 (points: None)
```
>>> assert get_hash(type(q5)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(gradation(70, 45, 0.15, 0.13, 0.1).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 45, 0.15, 0.11, 0.1).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 45, 0.15, 0.14, 0.005).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 45, 0.45, 0.3, 0.1).upper()) == '61e9c06ea9a85a5088a499df6458d276'
>>> assert get_hash(gradation(50, 25, 0.65, 0.3, 0.1).upper()) == '61e9c06ea9a85a5088a499df6458d276'
>>> assert get_hash(gradation(70, 35, 0.15, 0.13, 0.1).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 35, 0.45, 0.3, 0.1).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 35, 0.15, 0.11, 0.1).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 35, 0.15, 0.14, 0.005).upper()) == '44c29edb103a2872f519ad0c9a0fdaaa'
>>> assert get_hash(gradation(70, 35, 0.65, 0.3, 0.1).upper()) == '61e9c06ea9a85a5088a499df6458d276'
```


Tests for q6 (points: None)
```
>>> assert get_hash(type(q6)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(plasticity(10).upper()) == 'd20caec3b48a1eef164cb4ca81ba2587'
>>> assert get_hash(plasticity(30).upper()) == 'd20caec3b48a1eef164cb4ca81ba2587'
>>> assert get_hash(plasticity(45).upper()) == 'd20caec3b48a1eef164cb4ca81ba2587'
>>> assert get_hash(plasticity(55).upper()) == 'c1d9f50f86825a1a2302ec2449c17196'
>>> assert get_hash(plasticity(60).upper()) == 'c1d9f50f86825a1a2302ec2449c17196'
>>> assert get_hash(plasticity(70).upper()) == 'c1d9f50f86825a1a2302ec2449c17196'
```


Tests for q7 (points: None)
```
>>> assert get_hash(type(q7)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(silt_or_clay(55, 25).upper()) == '0d61f8370cad1d412f80b84d143e1257'
>>> assert get_hash(silt_or_clay(75, 25).upper()) == '0d61f8370cad1d412f80b84d143e1257'
>>> assert get_hash(silt_or_clay(30, 20).upper()) == '0d61f8370cad1d412f80b84d143e1257'
>>> assert get_hash(silt_or_clay(60, 30.7).upper()) == '0d61f8370cad1d412f80b84d143e1257'
>>> assert get_hash(silt_or_clay(55, 45).upper()) == '69691c7bdcc3ce6d5d8a1361f22d04ac'
>>> assert get_hash(silt_or_clay(75, 45).upper()) == '69691c7bdcc3ce6d5d8a1361f22d04ac'
>>> assert get_hash(silt_or_clay(30, 25).upper()) == '69691c7bdcc3ce6d5d8a1361f22d04ac'
>>> assert get_hash(silt_or_clay(60, 30.9).upper()) == '69691c7bdcc3ce6d5d8a1361f22d04ac'
```
```
>>> assert get_hash(silt_or_clay(20, 10).upper()) == '0d61f8370cad1d412f80b84d143e1257'
>>> assert get_hash(silt_or_clay(10, 7).upper()) == '69691c7bdcc3ce6d5d8a1361f22d04ac'
>>> assert get_hash(silt_or_clay(20, 17).upper()) == '69691c7bdcc3ce6d5d8a1361f22d04ac'
```
```
>>> assert get_hash(silt_or_clay(60, 30.8).upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
failure_message: If soil falls exactly on the "A"-line, is it clay or silt?


Tests for q8 (points: None)
```
>>> assert get_hash(type(q8)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(classify_fine(55, 25).upper()) == '1ee0bf89c5d1032317d13a2e022793c8'
>>> assert get_hash(classify_fine(75, 25).upper()) == '1ee0bf89c5d1032317d13a2e022793c8'
>>> assert get_hash(classify_fine(30, 20).upper()) == '5bc574a47246f122016869b32a6aa6f0'
>>> assert get_hash(classify_fine(60, 30.7).upper()) == '1ee0bf89c5d1032317d13a2e022793c8'
>>> assert get_hash(classify_fine(55, 45).upper()) == '002f27e5064e874ecf4f5def17d1b797'
>>> assert get_hash(classify_fine(75, 45).upper()) == '002f27e5064e874ecf4f5def17d1b797'
>>> assert get_hash(classify_fine(30, 25).upper()) == 'd01fd9b01e9dde8bd3dc247afbfb7218'
>>> assert get_hash(classify_fine(60, 30.9).upper()) == '002f27e5064e874ecf4f5def17d1b797'
```
```
>>> assert get_hash(classify_fine(20, 10).upper()) == '5bc574a47246f122016869b32a6aa6f0'
>>> assert get_hash(classify_fine(10, 7).upper()) == 'd01fd9b01e9dde8bd3dc247afbfb7218'
>>> assert get_hash(classify_fine(20, 17).upper()) == 'd01fd9b01e9dde8bd3dc247afbfb7218'
```
```
>>> assert get_hash(classify_fine(10, 5).upper()) == '6ee58e48c2cb2ff01807ccca4667e070'
>>> assert get_hash(classify_fine(20, 14).upper()) == '6ee58e48c2cb2ff01807ccca4667e070'
>>> assert get_hash(classify_fine(26, 20).upper()) == '6ee58e48c2cb2ff01807ccca4667e070'
```


Tests for q9 (points: None)
```
>>> assert get_hash(type(q9)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert classify_gravel(40, 4, 0.3, 0.14, 0.025, 0, 0).upper() == 'GW'
>>> assert classify_gravel(40, 4, 0.15, 0.13, 0.1, 0, 0).upper() == 'GP'
>>> assert classify_gravel(40, 10, 0.3, 0.14, 0.025, 30, 25).upper() == 'GW-GM'
>>> assert classify_gravel(40, 10, 0.3, 0.14, 0.025, 30, 15).upper() == 'GW-GC'
>>> assert classify_gravel(40, 10, 0.15, 0.13, 0.1, 30, 25).upper() == 'GP-GM'
>>> assert classify_gravel(40, 10, 0.15, 0.13, 0.1, 30, 15).upper() == 'GP-GC'
>>> assert classify_gravel(40, 20, 0.15, 0.13, 0.1, 30, 25).upper() == 'GM'
>>> assert classify_gravel(40, 20, 0.15, 0.13, 0.1, 30, 15).upper() == 'GC'
>>> assert classify_gravel(40, 20, 0.15, 0.13, 0.1, 25, 20).upper() == 'GC-GM'
```
```
>>> assert get_hash(classify_gravel(50, 1, 2, 1.1, 0.3, 0, 0).upper()) == 'c17d19f7520be36addbeb5d2c76ab7cf'
>>> assert get_hash(classify_gravel(50, 1, 2, 1.5, 0.3, 0, 0).upper()) == 'ad2d8ee7d788dcf41f399818f639cb64'
>>> assert get_hash(classify_gravel(50, 6, 0.45, 0.25, 0.1, 60, 40).upper()) == 'c31f61bf166e30ec55828503fad107d4'
>>> assert get_hash(classify_gravel(50, 6, 0.45, 0.25, 0.1, 60, 30).upper()) == 'acb50b1138e657e40d297035097d0bf1'
>>> assert get_hash(classify_gravel(50, 6, 0.45, 0.2, 0.1, 60, 40).upper()) == '03f9d3f62bd6324143e47935f642177a'
>>> assert get_hash(classify_gravel(50, 6, 0.45, 0.2, 0.1, 60, 30).upper()) == 'f9bb4d123b6b669ac79bd9465d619d00'
>>> assert get_hash(classify_gravel(50, 13, 0.15, 0.13, 0.1, 20, 17).upper()) == '64f3bd1741ab8d6ba545a1ae09bb8728'
>>> assert get_hash(classify_gravel(50, 13, 0.15, 0.13, 0.1, 55, 25).upper()) == '7359c3eb5c57547295a76ac1bf775b29'
>>> assert get_hash(classify_gravel(50, 13, 0.15, 0.13, 0.1, 10, 4).upper()) == '58a3dd33c44967a9f632495eaf1b41b0'
```
